---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Comparing two trading strategies: Technical vs. Fundamental analysis

```yaml
type: VideoExercise
key: 1750d82a05
xp: 50
```

`@projector_key`
4a3d2de89bb0656e8e954ca2a1347674
